# Welcome to an Intro to Python

## Preamble:

This interactive notebook is part of a series which can be found here:
https://github.com/DimitriosAstro/Astronomy  
The notebook is distributed by Dimitrios Theodorakis under the GNU General Public License v3.0.

Find more info at www.w3schools.com, www.python.org, www.learnpython.org  
The Numpy section was adapted from work by Ridwan Barbhuiyan (https://github.com/rbarbhuiyan)

The purpose of this Notebook is to get you ready to explore the other Notebooks which each explore a topic in Astronomy and Astrophysics using the python programing language.

## Start:

*Shift+Enter* on a code snippet to run the code. Most of the time you'll have to run the snippets in order or you'll get an error.

If you're stuck you can consult the docs at https://docs.python.org/3/, and https://numpy.org/doc/stable/.  
You can also get advice from www.stackoverflow.com.

**Comments** - All code should be commented for readability

```python
# This is a single line comment
'''This is a 
multiline comment'''
```

**Variables** - Assigning a variable. The dot says we want a float (decimal) not an integer. For b we assigned a string.

```python
a = 2
a = 2.
b = "Hello World!"
```

**Printing** - Print a variable (or anything else) using the print() fucntion.

```python
print(b)
```

**Maths** - Basic operators in python. Addition, subtraction, multiplication, division, and powers.

```python
c = 2+2
d = 2-2
e = 2*2
f = 2/2
g = 2**2
```

**Print Text and Numbers** - Use str() to convert the number to a string first and join the two strings together.

```python
print("My number is: "+str(c))
```

**Lists** - A list of numbers or strings or anything else!

```python
myList = [1,2,3,4,5]
print(myList)
```

Append to the list (at the end)

```python
myList.append(6)
print(myList)
```

**Tuples** - These are like lists but cannot be changed once created.

```python
myTuple = (1,2,3,4)
print(myTuple)
```

**Indexing** - Access an element in a list/tuple using its index. The first element is at index 0. You can also use negative indices and select multiple elements at once using a slice.

```python
print(myList[0])
print(myList[1])

print(myList[-1])
print(myList[-2])

print(myList[0:3])
```

**For Loops** - A for loop will do something (in this case print()) for each element (i) in the list/tuple.

```python
for i in myList:
    print(i)
```

**Numpy Library** - A library is a set of pre-made functions that you can call on in your code.  
The Numpy library is a very useful maths library. We can import it using the code below. The *as np* means we can just write *np* after importing instead of the full name *numpy*.

```python
import numpy as np
```

**Arrays** - A numpy array is generally faster than a list.

```python
array1 = np.arange(1,11)
print(array1)
```

**Dtypes** - Data types include int (integers), float (decimals), lists, tuples, and arrays.

```python
print(type(array1))
# notice the difference between that and the type for myList
print(type(myList))
```

**For Loop on Arrays**

```python
for i in array1:
    print(i*2) # times each element by 2
```

**Maths on an Array** - Operations can be performed on each element of an array.

```python
array2 = array1*2
print(array2)
```

**Linspace** - Another more controlled way of creating an array using a start, end, and number of points variables.

```python
start = 0
finish = 5
step = 0.5
number = int((finish - start) / step) + 1

array3 = np.linspace(start, finish, number)
print(array3)
```

**2-D Arrays** - You can create 2-D arrays by specifying a shape. Normally we initialise an array using ones or zeros then change the values inside.

```python
array4 = np.zeros((2,2))
array5 = np.ones((2,2))

print(array4)
print(array5)

# accessing a 2-D array [row, column], remember the first element is at index 0
print("-----")
print(array4[1,1])

# setting values in a 2-D array
print("-----")
array4[1,1] = 1
print(array4)
```

**Maths with Numpy** - Numpy has very useful maths functions. We can't list them all here but you can access the number pi and compute logs easily enough.

```python
print(np.pi)
print(np.log(10))
```

## Over to You

**Task 1:** Create a list with these 3 foods: apple, orange, and peach. Print the list.  
**Task 2:** Print just the third food in the list you just made.  
**Task 3:** Create an array using arange of the numbers 1-4. Print the square of those numbers.  
**Task 4:** Create a new 2-D array of zeros with size (3,3). Set the diagonals to 1 and print the array.

Solutions can be found in the non-interactive version of this document on the main repo (see very top).

<!-- grade_id: cell-e1aef7512be7ca92, points: 0, grade: false, solution: true -->
```python
# Task 1: Call your list fruit_list

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a14498ec61e8003d, points: 0, grade: false, solution: true -->
```python
# Task 2: Print the third food

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b9060d8b5ed8fbbb, points: 0, grade: false, solution: true -->
```python
# Task 3: Create an array of the numbers 1-5 using numpy's arange funstion
# Name your array my_array and print the square of your numbers

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9ac53fc5938da093, points: 0, grade: false, solution: true -->
```python
# Task 4: 3 by 3 array of zeros with diagonals as 1's - call this array Zeros

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9fdf474ddc426569, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(Zeros[0][0], 1.0)
assert_equal(Zeros[1][0], 0.0)
```

```python

```